# Assignment 1 — Problem 1

Implement a simple function `add(a, b)` that returns the sum of two numbers.

## Instructions (Giảng viên ghi chú)
- This notebook contains **one autograded problem** worth **10 points**.
- Students should **only edit** the cell marked `# YOUR CODE HERE`.
- Do **not** change function names or the cell structure.

<!-- grade_id: add_solution, points: 0, grade: false, solution: true -->
```python
def add(a, b):
    """Return the sum of a and b.

    Examples
    --------
    >>> add(2, 3)
    5
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: add_tests, points: 10, grade: true, solution: false -->
```python
# AUTOGRADED TESTS (Do not modify!)
# If your implementation is correct, these assertions will pass.

assert callable(add), "add must be callable"
assert add(0, 0) == 0
assert add(2, 3) == 5
assert add(-4, 9) == 5
assert add(1.5, 2.5) == 4.0

# Additional randomized checks
import random
for _ in range(50):
    x = random.randint(-1000, 1000)
    y = random.randint(-1000, 1000)
    assert add(x, y) == x + y

print("All tests passed!")

```